<a href="https://colab.research.google.com/github/hosafxd/LavaAlgorithm/blob/main/qualification_(1).ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# RandomLogistic
In this assigment you will implement RandomLogistic Regression.
- First you will read the image using the given instructions.
- Then you will flatten and normalize the image.
- After you implement model compute function will be used to get the result.

## **IMPORTANT WARNINGS**
- **DO NOT USE** any additional libraries other than the standard Python libraries and the ones in the requirements.txt
- **DO NOT DELETE** any cells even if it they are empty
- While implementing the functions you should only change the cells with the statement #YOUR CODE HERE \
- If you change other than those blocks your notebook might not be graded properly

```python
import numpy as np
from PIL import Image
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

## Image reading
Read the image using PIL library and convert the image to a numpy array using 'np.asarray' function

<!-- grade_id: cell-a6681d0d24df745a, points: 0, grade: false, solution: true -->
```python
def read_image(img_path):
    # YOUR CODE HERE
    img = Image.open(img_path)
    img = np.asanyarray(img)


   # raise NotImplementedError()
    return img
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

```python
img = read_image('inzva.jpeg')
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-27634dfa545661ce, points: 10, grade: true, solution: false -->
```python
assert type(img) == np.ndarray
assert img.shape == (225, 225, 3)
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

## Flatten the Image
The given image has the shape (1000, 1000, 3). We want you to reshape the image into (3000000,)

<!-- grade_id: cell-52062b23598241d2, points: 0, grade: false, solution: true -->
```python
def flatten_img(img):
    # YOUR CODE HERE
    img_flatten = img.reshape(-1)


    #raise NotImplementedError()
    return img_flatten
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

```python
img_flatten = flatten_img(img)
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-7b2b4622d95c08a9, points: 10, grade: true, solution: false -->
```python
assert img_flatten.shape == (151875,) or img_flatten.shape == (151875,1)
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

## Normalize the Image
You should implement the following formula to normalize the image:
$$ x = \frac{x - min(x)}{max(x)-min(x)} $$

<!-- grade_id: cell-9b4666d8308a183b, points: 0, grade: false, solution: true -->
```python
def normalize_img(img):
    # YOUR CODE HERE
    normalized_img = (img - np.min(img)) / (np.max(img) - np.min(img))


    # raise NotImplementedError()
    return normalized_img
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

```python
img_normalized = normalize_img(img_flatten)
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-61a350c6477164bb, points: 10, grade: true, solution: false -->
```python
assert np.max(img_normalized) == 1.0 and np.min(img_normalized) == 0.0
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

## Random Logistic
Set weights using 'np.random.randn' and set bias to zero. Be careful with the dimensions. Then implement the sigmoid function ($\sigma $) using the following formula:
$$\sigma(x) = \frac{1}{ 1 + \exp(-x)} $$
Then implement the forward pass by:
$$y = \sigma((W * x) + b)$$

<!-- grade_id: cell-7ef2b26abd49de09, points: 0, grade: false, solution: true -->
```python
class RandomLogistic:
    def __init__(self, input_shape):
        np.random.seed(2024)
        ## Define your weights and bias below.
        # YOUR CODE HERE
        self.weights = np.random.randn(*input_shape) # Unpack the tuple
        self.bias = 0.0


       # raise NotImplementedError()

    def sigmoid(self, x):
        # YOUR CODE HERE
        return 1 / (1 + np.exp(-x))


      #  raise NotImplementedError()

    def compute(self, x):
        # YOUR CODE HERE
        linear_combination = np.dot(self.weights, x) + self.bias
        y = self.sigmoid(linear_combination)


      #  raise NotImplementedError()
        return y
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

```python
random = RandomLogistic(input_shape=img_flatten.shape)
y = random.compute(img_normalized)
z = np.load('result.npy')
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-3a57a78b3ebce3d2, points: 20, grade: true, solution: false -->
```python
assert np.sum(np.abs(y-z)) < 1e-9
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

# Edit Distance

- Complete the function below to find the edit distance between two strings.
- Edit distance between two words is defined as the minimum number of operations needed to transform one word into another, in which the only operations allowed are inserting, deleting or replacing a character.
(insertion: inzva → inz**i**va, deletion: inz**v**a → inza, replacement: inzva → in**o**va)

**For this question, use Python’s standard libraries only.**\
Make sure that your algorithm works correctly as there are also hidden test cases that are not shown here.

<!-- grade_id: cell-3f5cca84206b51d5, points: 0, grade: false, solution: true -->
```python
def edit_distance(word1, word2):
    # YOUR CODE HERE
    memo = {}

    def dp(i, j):

        if (i, j) in memo:
            return memo[(i, j)]


        if i == 0:
            return j

        if j == 0:
            return i

        if word1[i - 1] == word2[j - 1]:
            memo[(i, j)] = dp(i - 1, j - 1)
        else:

            memo[(i, j)] = min(
                dp(i - 1, j) + 1,  # Deletion
                dp(i, j - 1) + 1,  # Insertion
                dp(i - 1, j - 1) + 1  # Replacement
            )
        return memo[(i, j)]

    return dp(len(word1), len(word2))

#    raise NotImplementedError()
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-e68c3f781e28fcf7, points: 6, grade: true, solution: false -->
```python
assert edit_distance("inzva", "inziva") == 1
assert edit_distance("inzva", "inova") == 1
assert edit_distance("inzva", "dlsg") == 5
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-68a823ce760d1540, points: 15, grade: true, solution: false -->
```python
### DO NOT DELETE THIS CELL
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```

<!-- grade_id: cell-f2b3f99ef05313d8, points: 29, grade: true, solution: false -->
```python
### DO NOT DELETE THIS CELL
```

Output:
```
Running cells with 'Python 3.11.9' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/FurkanÖztürk/AppData/Local/Microsoft/WindowsApps/python3.11.exe -m pip install ipykernel -U --user --force-reinstall'
```